**We're going to help the guards out and compute whether a collection of swallows can carry a collection of coconuts between them. We know that "a five ounce bird could not carry a one pound coconut," so we will assume that a 5.5 ounce bird can carry a one pound coconut. More specifically, we will assume that every 5.5 ounces of bird can carry one pound of coconut.**

Prompt the user for the number of ounces of bird 
by asking: "How many ounces of birds are carrying the coconuts? ".
Prompt the user for the number of pounds of coconut by asking: 
"How many pounds of coconuts are there? ".
Remember to convert these values to floats! 
You can assume that the user input is formatted correctly.

<!-- grade_id: cell-cf9726a45fb25ba2, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
# def carrying_coconuts(ounces_of_birds, pounds_of_coconuts):

### BEGIN SOLUTION
def carrying_coconuts(ounces_of_birds, pounds_of_coconuts):
    ounces_of_birds = float(input('How many ounces of birds are carrying the coconuts?'))
    pounds_of_coconuts = float(input('How many pounds of coconuts are there?'))
    
    return "Yes! Carrying the coconuts is possible." if ounces_of_birds / pounds_of_coconuts >= 5.5 else "No. Carrying the coconuts is impossible."
### END SOLUTION


```

# **Sample run**

How many ounces of birds are carrying the coconuts? 5
How many pounds of coconuts are there? 1
No. Carrying the coconuts is impossible.


How many ounces of birds are carrying the coconuts? 6.2
How many pounds of coconuts are there? 1.1
Yes! Carrying the coconuts is possie

py
How many ounces of birds are carrying the coconuts? 17
How many pounds of coconuts are there? 3
Yes! Carrying the cos is conutpossi
ts.py
How many ounces of birds are carrying the coconuts? 12.5
How many pounds of coconuts are there? 2.5
No. Carrying the coconuts is impossible.


<!-- grade_id: cell-308136d4df2f8427, points: 1, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS

import unittest

class TestCarryingCoconuts(unittest.TestCase):

    def test_carrying_coconuts_possible(self):
        # expected_output = "Yes! Carrying the coconuts is possible."
        result = carrying_coconuts(10, 2)
        self.assertEqual(result, "Yes! Carrying the coconuts is possible.")

    def test_carrying_coconuts_impossible(self):
        result = carrying_coconuts(5, 2)
        self.assertEqual(result, "No. Carrying the coconuts is impossible.")

    def test_carrying_coconuts_edge_case(self):
        # Testing with the minimum allowed values
        result = carrying_coconuts(0, 1)
        self.assertEqual(result, "No. Carrying the coconuts is impossible.")

    def test_carrying_coconuts_large_values(self):
        result = carrying_coconuts(100, 20)
        self.assertEqual(result, "Yes! Carrying the coconuts is possible.")

    def test_carrying_coconuts_equal_values(self):
        result = carrying_coconuts(5.5, 1)
        self.assertEqual(result, "Yes! Carrying the coconuts is possible.")

### END HIDDEN TESTS
```

Make a list of cheeses containing "Muenster", "Cheddar", and "Red Leicester"; these are the cheeses you have in your shop.

Begin by printing the string "Good morning. Welcome to the National Cheese Emporium!" to the console.

Next, in a loop, we will repeatedly ask the user which cheese they would like to buy. If the user enters the exact name of a cheese that you have, affirm that you have the cheese they asked for in the format "We have {}, yessir.". If the user does not enter the exact name of a cheese that you have, say that you don't have that cheese in the form "I'm afraid we don't have any {}." and then reprompt the user to ask for another cheese.

The user is also allowed to enter special questions. If the user enters either of the strings "You... do have some cheese, don't you?" or "Have you in fact got any cheese here at all?", you must reply by listing the number of cheeses that you have in the format "We have {} cheese(s)!", along with the name of each type of cheese you have, one on each

**For anything that is not specified in the above specification, your program can behave in anyway you'd like. For example, you can customize the prompts and messages.** line.

<!-- grade_id: cell-274ac8afe3538dae, points: 0, grade: false, solution: true -->
```python
# YOUR SOLUTION HERE

### BEGIN SOLUTION

cheeses = ["Muenster", "Cheddar", "Red Leicester"]

# welcome message
print("Good morning. Welcome to the National Cheese Emporium!")

# Loop to repeatedly ask the user 
while True:
    user_input = input("Which cheese would you like to buy? ")

    # special questions
    if user_input == "You... do have some cheese, don't you?" or user_input == "Have you in fact got any cheese here at all?":
        print("We have {} cheese(s)!".format(len(cheeses)))
        for cheese in cheeses:
            print(cheese)
        break

    # check if the entered cheese is available
    elif user_input in cheeses:
        print(f"We have {user_input}, yessir.")
        break

    else:
        print(f"I'm afraid we don't have any {user_input}.")

### END SOLUTION

```

Output:
```
Good morning. Welcome to the National Cheese Emporium!
Which cheese would you like to buy?  1
I'm afraid we don't have any 1.
Which cheese would you like to buy?  1
I'm afraid we don't have any 1.
Which cheese would you like to buy?  Cheddar
We have Cheddar, yessir.
```

**Sample Run**

Good morning. Welcome to the National Cheese Emporium!
What would you like? Red Windsor
I'm afraid we don't have any Red Windsor.
What would you like? Lancashire
I'm afraid we don't have any Lancashire.
What would you like? cheddar
I'm afraid we don't have any cheddar.
What would you like? Cheddar
We have Cheddar, yessir.

Good morning. Welcome to the National Cheese Emporium!
What would you like? Red Windsor
I'm afraid we don't have any Red Windsor.
What would you like? cHeDdAr
I'm afraid we don't have any cHeDdAr.
What would you like? Have you in fact got any cheese here at all?
We have 3 cheese(s)!
Muenster
Cheddar
Red Leicester
What would you like? exit
I'm afraid we don't have any exit.
What would you like? LET ME OUT
I'm afraid we don't have any LET ME OUT.
What would you like? Cheddar
We have Cheddar, yessir.

<!-- grade_id: cell-5057de5b041f6f13, points: 0, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS

import unittest


class TestingCheese(unittest.TestCase):
    def 



### END HIDDEN TESTS
```

Output:
```
EEE
======================================================================
ERROR: test_cheese_selection_existing (__main__.TestCheeseShopScript)
----------------------------------------------------------------------
Traceback (most recent call last):
  File "C:\Users\User\AppData\Local\Programs\Python\Python310\lib\unittest\mock.py", line 1369, in patched
    return func(*newargs, **newkeywargs)
  File "C:\Users\User\AppData\Local\Temp\ipykernel_13976\2971552085.py", line 14, in test_cheese_selection_existing
    carrying_coconuts()
  File "C:\Users\User\AppData\Local\Temp\ipykernel_13976\18855141.py", line 10, in carrying_coconuts
    ounces_of_birds = float(input('How many ounces of birds are carrying the coconuts?'))
ValueError: could not convert string to float: 'Cheddar'

======================================================================
ERROR: test_cheese_selection_non_existing (__main__.TestCheeseShopScript)
-------------------------------------------------------------------
```

```python

```

```python

```

```python

```

```python

```